### Codio Activity 4.2: Complex Joins on Datasets

**Expected Time: 60 Minutes**

**Total Points: 10**

![](images/kiva.png)

This assignment focuses on using pandas merge to answer questions using multiple data sources.  Here you will combine data and use many of the earlier `pandas` methods together with the `merge` function to extract insights from our data source.  The data comes from the Kiva loan provider organization and was part of a kaggle Data Science for good contest [here](https://www.kaggle.com/kiva/data-science-for-good-kiva-crowdfunding).

#### Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)
- [Problem 5](#Problem-5)

### The Data

Kiva allows users to fund loans to small organizations around the world.  The four datasets contain a subset of the data provided on kaggle and describe information on the individual loan, when it was given, where the borrowers were located, and what the loans were being used for.  Below, the data is loaded into four DataFrames and the first two rows of each displayed.  Note the shared `id` column across all four datasets.

```python
import pandas as pd
```

```python
crowdsource = pd.read_csv('data/kiva/crowdsource.csv')
crowdsource.head(2)
```

Output:
```
id                posted_time                funded_time  lender_count
0   658776  2014-01-17 21:21:10+00:00  2014-02-05 17:57:55+00:00            33
1  1314847  2017-06-07 02:02:41+00:00  2017-06-21 17:10:38+00:00             9
```

```python
demographics = pd.read_csv('data/kiva/demographics.csv')
demographics.head(2)
```

Output:
```
id      country                 region borrower_genders
0   658776  El Salvador      Ciudad El Triunfo             male
1  1314847  Philippines  Bais, Negros Oriental           female
```

```python
financials = pd.read_csv('data/kiva/financials.csv')
financials.head(2)
```

Output:
```
id  funded_amount currency  term_in_months
0   658776         1000.0      USD            20.0
1  1314847          225.0      PHP            13.0
```

```python
use = pd.read_csv('data/kiva/use.csv')
use.head(2)
```

Output:
```
id activity        sector  \
0   658776  Vehicle  Personal Use   
1  1314847     Pigs   Agriculture   

                                                 use  
0  to purchase a motorcycle in order to travel fr...  
1  to buy feed and other supplies like vitamins t...
```

[Back to top](#Index:) 

### Problem 1

#### Kenyan Loan Amounts

**5 Points**

Use the `demographics` and `financials` data to determine the average loan amount for the country Kenya.  Save your response as a float to `ans1` below.

<!-- grade_id: cell-30ba4362a6bbb79d, points: 0, grade: false, solution: true -->
```python
### GRADED

ans1 = None

### BEGIN SOLUTION
ans1 = pd.merge(left=demographics, right=financials, on='id').query('country == "Kenya"')[['funded_amount']].mean().values[0]
type(ans1)
### END SOLUTION

# Answer check
print(ans1)
print(type(ans1))
```

Output:
```
443.4723441615452
<class 'numpy.float64'>
```

<!-- grade_id: cell-3c3e3fda1e1b70e8, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
crowdsource_ = pd.read_csv('data/kiva/crowdsource.csv')
demographics_ = pd.read_csv('data/kiva/demographics.csv')
financials_ = pd.read_csv('data/kiva/financials.csv')
use_ = pd.read_csv('data/kiva/use.csv')
ans1_ = pd.merge(demographics_.loc[demographics_['country'] == 'Kenya'], financials_, on = 'id')[['funded_amount']].mean().values[0]
#
#
#
assert type(ans1_) == type(ans1)
assert ans1_ == ans1
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

### El Salvador Top Sector

**5 Points**

Use the `demographics` and `use` datasets to determine the sector of work that received the most loans in El Salvador.  Assign your result as a string to `ans2` below.

<!-- grade_id: cell-4037ebf501493e72, points: 0, grade: false, solution: true -->
```python
### GRADED

ans2 = None

### BEGIN SOLUTION
ans2 = pd.merge(left=demographics, right=use, on='id').query('country == "El Salvador"')['sector'].value_counts().index[0]
type(ans2)
### END SOLUTION

# Answer check
print(ans2)
print(type(ans2))
```

Output:
```
Agriculture
<class 'str'>
```

<!-- grade_id: cell-2bf2a356d27da211, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
crowdsource_ = pd.read_csv('data/kiva/crowdsource.csv')
demographics_ = pd.read_csv('data/kiva/demographics.csv')
financials_ = pd.read_csv('data/kiva/financials.csv')
use_ = pd.read_csv('data/kiva/use.csv')
ans2_ = pd.merge(demographics_.loc[demographics_['country'] == 'El Salvador'], use_, on = 'id')['sector'].value_counts().index[0]
#
#
#
assert type(ans2_) == type(ans2)
assert ans2_ == ans2
### END HIDDEN TESTS
```

### Problem 3

#### Pakistan Loans in Agriculture

**5 Points**

Combine the `use`, `demographics`, and `financials` data to determine the total amount of loans for Agriculture in Pakistan.  Assign your results as a float to `ans3` below.

<!-- grade_id: cell-00a00bf68209058f, points: 0, grade: false, solution: true -->
```python
### GRADED

ans3 = None

### BEGIN SOLUTION
p1 = pd.merge(use, demographics, on = 'id')
a = pd.merge(p1, financials, on = 'id')
b = a.loc[a['country'] == 'Pakistan']
ans3 = b.loc[b['sector'] == 'Agriculture'][['funded_amount']].sum().values[0]

### END SOLUTION

# Answer check
print(ans3)
print(type(ans3))
```

Output:
```
30900.0
<class 'numpy.float64'>
```

<!-- grade_id: cell-79780abf2a643093, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
crowdsource_ = pd.read_csv('data/kiva/crowdsource.csv')
demographics_ = pd.read_csv('data/kiva/demographics.csv')
financials_ = pd.read_csv('data/kiva/financials.csv')
use_ = pd.read_csv('data/kiva/use.csv')
p1_ = pd.merge(use_, demographics_, on = 'id')
a_ = pd.merge(p1_, financials_, on = 'id')
b_ = a_.loc[a_['country'] == 'Pakistan']
ans3_ = b_.loc[b_['sector'] == 'Agriculture'][['funded_amount']].sum().values[0]
#
#
assert type(ans3_) == type(ans3)
assert ans3_ == ans3
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4

#### Top Total Loan Sector

**5 Points**


What sector received the most total dollars in funding?  Assign your response as a string to `ans4` below.

<!-- grade_id: cell-a8d7b30a738b7adc, points: 0, grade: false, solution: true -->
```python
### GRADED

ans4 = None

### BEGIN SOLUTION
ans4 = pd.merge(financials, use, on = 'id').groupby('sector')[['funded_amount']].sum().sort_values(by = 'funded_amount', ascending = False).index[0]
### END SOLUTION

# Answer check
print(ans4)
print(type(ans4))
```

Output:
```
Agriculture
<class 'str'>
```

<!-- grade_id: cell-c4a056a912aea7e0, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
crowdsource_ = pd.read_csv('data/kiva/crowdsource.csv')
demographics_ = pd.read_csv('data/kiva/demographics.csv')
financials_ = pd.read_csv('data/kiva/financials.csv')
use_ = pd.read_csv('data/kiva/use.csv')
ans4_ = pd.merge(financials_, use_, on = 'id').groupby('sector')[['funded_amount']].sum().sort_values(by = 'funded_amount', ascending = False).index[0]
#
#
assert type(ans4_) == type(ans4)
assert ans4_ == ans4
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 5

#### Top Loan by Lender Amount

**5 Points**


Note that each loan is a combination of lenders.  Determine which loan sector has the highest ratio of dollar to lender amount.  Assign the sector with the highest ratio as a string to `ans5` below. 

<!-- grade_id: cell-1739efe5ab750822, points: 0, grade: false, solution: true -->
```python
### GRADED

ans5 = None

# YOUR CODE HERE
temp = pd.merge(left=use, right=financials, on='id', how='left')
temp = pd.merge(left=temp, right=crowdsource, on='id')
temp_grouped = temp[['sector', 'funded_amount', 'lender_count']].groupby('sector').sum()
temp_grouped['ratio'] = temp_grouped['funded_amount']/temp_grouped['lender_count']
temp_grouped = temp_grouped.sort_values(by='ratio', ascending=False)
ans5 = temp_grouped.index[0]

# Answer check
print(ans5)
print(type(ans5))


```

Output:
```
Clothing
<class 'str'>
```

<!-- grade_id: cell-95ae34d4661d4478, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
crowdsource_ = pd.read_csv('data/kiva/crowdsource.csv')
demographics_ = pd.read_csv('data/kiva/demographics.csv')
financials_ = pd.read_csv('data/kiva/financials.csv')
use_ = pd.read_csv('data/kiva/use.csv')
a_ = pd.merge(financials_, use_, on = 'id')
b_ = pd.merge(a_, crowdsource, on = 'id')
b_['ratio'] = b_['funded_amount']/b_['lender_count']
print(b_)
ans5_ = b_.groupby('sector')['ratio'].max().idxmax()
#
#
#
assert type(ans5_) == type(ans5)
assert ans5_ == ans5
### END HIDDEN TESTS
```

Output:
```
id  funded_amount currency  term_in_months  \
0      658776         1000.0      USD            20.0   
1     1314847          225.0      PHP            13.0   
2      863063         1150.0      PEN             6.0   
3     1184347         1700.0      AMD            26.0   
4      729745          400.0      UGX             8.0   
...       ...            ...      ...             ...   
9995   679499          400.0      PKR            12.0   
9996   873525          375.0      KES            14.0   
9997   917686         1375.0      XOF             8.0   
9998   905789          450.0      PHP            13.0   
9999  1216411          125.0      PHP            16.0   

                       activity        sector  \
0                       Vehicle  Personal Use   
1                          Pigs   Agriculture   
2                     Bookstore        Retail   
3                   Photography      Services   
4                 Fuel/Firewood        Retail   
...                         ... 
```

```python

```

```python

```