### Codio Activity 4.5: String Operations


**Expected Time: 60 Minutes**

**Total Points: 20**


For this activity, you will focus on using string methods on `pandas` series.  Following the examples from video 4.8, a wikipedia page on the states of the russian federation will be the dataset.  Below is a map of the boundaries for the states. 

#### Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)

![](https://upload.wikimedia.org/wikipedia/commons/thumb/7/7d/Map_of_federal_subjects_of_Russia_%282014%29.svg/1380px-Map_of_federal_subjects_of_Russia_%282014%29.svg.png)

### Getting the Data

Below, we access the data directly using the url and the `read_html` method from `pandas`.  This method reads any table from a web url in as a list of dataframes.  The data we are interested in is located in the fifth table on the page so we index the list of data accordingly below.

```python
import pandas as pd
```

```python
russian_states = pd.read_csv('data/russian_states.csv', index_col = 0)
```

```python
russian_states.head()
```

Output:
```
Code                        Name Capital/Administrative centre[a]  Flag  \
0     1         Adygea, Republic of                           Maykop   NaN   
1     2  Bashkortostan, Republic of                              Ufa   NaN   
2     3       Buryatia, Republic of                         Ulan-Ude   NaN   
3     4              Altai Republic                    Gorno-Altaysk   NaN   
4     5       Dagestan, Republic of                      Makhachkala   NaN   

   Coatof arms      Type Federal district Economic region Area(km2)[16]  \
0          NaN  republic         Southern  North Caucasus          7600   
1          NaN  republic            Volga            Ural        143600   
2          NaN  republic      Far Eastern   East Siberian        351300   
3          NaN  republic         Siberian   West Siberian         92600   
4          NaN  republic  North Caucasian  North Caucasus         50300   

  Population[17]  Populationdensity(/km2)  Yearestablished  
0         447109      
```

[Back to top](#Index:) 

### Problem 1

#### Using the `contains` method

**5 Points**

Use the string method `contains` to subset the data based on entries in the Economic Region column containing `Siberian`.  Assign your response as a DataFrame to `ans1` below.

<!-- grade_id: cell-69b100fd1408e91a, points: 0, grade: false, solution: true -->
```python
### GRADED

ans1 = None

### BEGIN SOLUTION
ans1 = russian_states[russian_states['Economic region'].str.contains('Siberian')]
type(ans1)
### END SOLUTION

# Answer check
print(ans1.shape)
ans1.head()
```

Output:
```
(15, 12)
    Code                    Name Capital/Administrative centre[a]  Flag  \
2      3   Buryatia, Republic of                         Ulan-Ude   NaN   
3      4          Altai Republic                    Gorno-Altaysk   NaN   
16    17           Tuva Republic                            Kyzyl   NaN   
18    19  Khakassia, Republic of                           Abakan   NaN   
21    22              Altai Krai                          Barnaul   NaN   

    Coatof arms      Type Federal district Economic region Area(km2)[16]  \
2           NaN  republic      Far Eastern   East Siberian        351300   
3           NaN  republic         Siberian   West Siberian         92600   
16          NaN  republic         Siberian   East Siberian        170500   
18          NaN  republic         Siberian   East Siberian         61900   
21          NaN      krai         Siberian   West Siberian        169100   

   Population[17]  Populationdensity(/km2)  Yearestablished  
2          981238    
```

<!-- grade_id: cell-1de07c86f4ae43cc, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
russian_states_ = pd.read_csv('data/russian_states.csv', index_col = 0)
ans1_ = russian_states_.loc[russian_states_['Economic region'].str.contains("Siberian")]
#
#
#
assert type(ans1) == type(ans1_)
pd.testing.assert_frame_equal(ans1, ans1_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 2

#### Using the `startswith` method

**5 Points**

Subset the data based on entries in the Economic region column that start with `North`.  Assign your answer as a DataFrame to `ans2` below.

<!-- grade_id: cell-0f898f719d53e4c3, points: 0, grade: false, solution: true -->
```python
### GRADED

ans2 = None

### BEGIN SOLUTION
ans2 = russian_states[russian_states['Economic region'].str.startswith('North')]
type(ans2)
### END SOLUTION

# Answer check
print(ans2.shape)
ans2.head()
```

Output:
```
(22, 12)
   Code                        Name Capital/Administrative centre[a]  Flag  \
0     1         Adygea, Republic of                           Maykop   NaN   
4     5       Dagestan, Republic of                      Makhachkala   NaN   
5     6     Ingushetia, Republic of      Magas(Largest city: Nazran)   NaN   
6     7   Kabardino-Balkar Republic                          Nalchik   NaN   
8     9  Karachay-Cherkess Republic                        Cherkessk   NaN   

   Coatof arms      Type Federal district Economic region Area(km2)[16]  \
0          NaN  republic         Southern  North Caucasus          7600   
4          NaN  republic  North Caucasian  North Caucasus         50300   
5          NaN  republic  North Caucasian  North Caucasus          4000   
6          NaN  republic  North Caucasian  North Caucasus         12500   
8          NaN  republic  North Caucasian  North Caucasus         14100   

  Population[17]  Populationdensity(/km2)  Yearestablished  
0         
```

<!-- grade_id: cell-99d4a7e77d5acc3b, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
russian_states_ = pd.read_csv('data/russian_states.csv', index_col = 0)
ans2_ = russian_states_.loc[russian_states_['Economic region'].str.startswith('North')]
#
#
#
assert type(ans2) == type(ans2_)
pd.testing.assert_frame_equal(ans2, ans2_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 3: Using the `upper` method

**5 Points**

Use the `upper` method to create a series with the Federal District column all uppercased.  Assign your response as a Series to `ans3` below.

<!-- grade_id: cell-228226682431b099, points: 0, grade: false, solution: true -->
```python
### GRADED

ans3 = None

### BEGIN SOLUTION
ans3 = russian_states['Federal district'].str.upper()
type(ans3)
### END SOLUTION

# Answer check
print(type(ans3))
ans3.head()
```

Output:
```
<class 'pandas.core.series.Series'>
0           SOUTHERN
1              VOLGA
2        FAR EASTERN
3           SIBERIAN
4    NORTH CAUCASIAN
Name: Federal district, dtype: object
```

<!-- grade_id: cell-513335067d60d99a, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
russian_states_ = pd.read_csv('data/russian_states.csv', index_col = 0)
ans3_ = russian_states_['Federal district'].str.upper()
#
#
#
assert type(ans3) == type(ans3_)
pd.testing.assert_series_equal(ans3, ans3_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

### Problem 4

#### Examining the Population

**5 Points**
 

Much like the example in the videos, the Population column contains problematic characters that need to be replaced before the column can be converted to a float datatype.  Replace these values and convert the Population column to float datatypes.  Assign the your response as a series to `ans4` below.  

**Note:** Create a new dataframe for your cleaned data. Don’t modify the original russian_states dataframe.

<!-- grade_id: cell-74ef2f0c657d0393, points: 0, grade: false, solution: true -->
```python
### GRADED

ans4 = None

### BEGIN SOLUTION
ans4 = russian_states['Population[17]'].str.replace(",","").str.replace("[22]","").str.replace("[23]","").astype(float)

type(ans4)
### END SOLUTION

# Answer check
print(type(ans4))
ans4.head()
```

Output:
```
<class 'pandas.core.series.Series'>
0     447109.0
1    4104336.0
2     981238.0
3     202947.0
4    2576531.0
Name: Population[17], dtype: float64
```

<!-- grade_id: cell-e5abae9366ae5429, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
russian_states_ = pd.read_csv('data/russian_states.csv', index_col = 0)
ans4_ = russian_states_['Population[17]'].str.replace("\[22\]", " ").str.replace("\[23\]", " ").str.replace(',', '').astype('float')

#
#
#
assert type(ans4) == type(ans4_)
pd.testing.assert_series_equal(ans4, ans4_)
### END HIDDEN TESTS
```

Output:
```
File /opt/anaconda3/lib/python3.11/site-packages/pandas/core/dtypes/astype.py:183, in astype_array(values, dtype, copy)
    180     values = values.astype(dtype, copy=copy)
    182 else:
--> 183     values = _astype_nansafe(values, dtype, copy=copy)
    185 # in pandas we don't store numpy str dtypes, so convert to object
    186 if isinstance(dtype, np.dtype) and issubclass(values.dtype.type, str):

File /opt/anaconda3/lib/python3.11/site-packages/pandas/core/dtypes/astype.py:134, in _astype_nansafe(arr, dtype, copy, skipna)
    130     raise ValueError(msg)
    132 if copy or arr.dtype == object or dtype == object:
    133     # Explicit copy, or required since NumPy can't view from / to object.
--> 134     return arr.astype(dtype, copy=True)
    136 return arr.astype(dtype, copy=copy)

ValueError: could not convert string to float: '1966801[22]'
```

```python

```